---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Il s'agit du sujet de l'interrogation d'étude de cas, merci de répondre à chaque question

`@instructions`
Pour la matrice M définie plus bas donner la syntaxe nécessaire à sélectionner la première colonne

`@hint`


`@pre_exercise_code`
```{r}
install.packages("gapminder")
```

`@sample_code`
```{r}
library(gapminder)
gapminder


```

`@solution`
```{r}

```

`@sct`
```{r}

```
